# Assignment 1: Getting Started

This assignment covers Variables, Operators & Conditionals.

This assignment is out of 10 points, worth 10% of your grade.

## How to complete assignments

Whenever you see:

```
# YOUR CODE HERE
raise NotImplementedError()
```

You need to replace this section with some code that answers the questions and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with 'assert' statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

## Tips & Tricks

The following are a couple tips & tricks that may help you if you get stuck on anything.

### Printing Variables

A reminder that you can (and should) print and check variables as you go.

This allows you to check what values they hold, and fix things if anything unexpected happens.

```python
# Define a variable
math_result = 2 * 4

# Print out the value(s) of a variable.
print(math_result)
```

### Restarting the Kernel

A reminder that sometimes if you run cells out of order, you can end up over-writing things in your namespace. 

If things seem to go weird, a good first step is to restart the kernel, which you can do from the kernel menu above.

Even if everything seems to be working, it's a nice check to 'Restart & Run All', to make sure everything runs properly in order.

## Example Question

Define a variable called `my_int` that stores the value 17. 

```python
# YOUR CODE HERE
my_int = 17
```

```python
# This is a test cell. It checks that the code you wrote does what it's supposed to
assert  my_int                     # This tests that a variable exists
assert  isinstance(my_int, int)    # This tests that a variable is of a particular type
assert  my_int == 17               # This tests that a variable has a specified value
```

## SETUP

Run the following cell before you start working on the assignment. 

This is going to make a helper folder on your computer, for some of the grading and checking things in this assignment. 

You do not need to interact with this folder maunally at all - everything is automatic.

```python
import os
if not os.path.exists('A1Code'):
    os.mkdir('A1Code')
```

## Part I - Defining Variables

First, we will start by defining some Python variables. 

### Q1 - Create Variables (1 point)

Create the following variables:
- a variable called `number` that stores the value 23
- a variable called `decimal` that stores the value 12.5
- a variable called `word` that stores the value 'cogs18'
- a variable called `truth` that stores the value True

<!-- grade_id: cell-aba100abfe23c75f, points: 0, grade: false, solution: true -->
```python
# Create the variables

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6d742eec19ccc8b1, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q1

# Check that variable `number` is an int, and has the correct value
assert isinstance(number, int)
assert number == 23
```

<!-- grade_id: cell-4757efea652f25f7, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q1

# Check that variable `decimal` is a float, and has the correct value
assert isinstance(decimal, float)
assert decimal == 12.5
```

<!-- grade_id: cell-8b0f7e2f80328a33, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q1

# Check that variable `words` is a string, and has the correct value
assert isinstance(word, str)
assert word == "cogs18"
```

<!-- grade_id: cell-bea89cf2bc700120, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q1

# Check that variable `truth` is a boolean, and has the correct value
assert isinstance(truth, bool)
assert truth == True
```

## Part II - Operations

### Q2 - Python as a Calculator (0.75 points)

Do the following calculations:
- Set `ans_1` as the result of int_1 multiplied by int_2
- Set `ans_2` as the result of float_1 divided by float_2, with int_2 then added
- Set `ans_3` as the result of the remainder of int_1 divided by int_2

```python
# These variables provided for you
int_1 = 47
int_2 = 12

float_1 = 47.5
float_2 = 19.0
```

<!-- grade_id: cell-470235f72bd8d68d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-12af4e4d3179f69e, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q2  [Note: Includes Hidden Tests]

assert isinstance(ans_1, int)

```

<!-- grade_id: cell-aeeac9dbb5198fef, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q2  [Note: Includes Hidden Tests]

assert isinstance(ans_2, float)

```

<!-- grade_id: cell-4bdf00498b0db006, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q2  [Note: Includes Hidden Tests]

assert isinstance(ans_3, int)

```

### Q3 - String Manipulations (0.75 points)

Using the strings provided, concatenate them together into the string `Python is fun!`. 

Save this result to a variable called `str_output`. 

Note: make sure to add spaces between words when concatenating together. 

```python
# These variables provided for you
str_1 = "Python"
str_2 = "is"
str_3 = "fun!"
```

<!-- grade_id: cell-eff2d0b53b5430eb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b9a1407c39acc29e, points: 0.75, grade: true, solution: false -->
```python
## TESTS FOR Q3

assert isinstance(str_output, str)
assert str_output == "Python is fun!"
```

### Q4 - Boolean Comparisons (1 point)

Using the variables provided, do the following boolean logical comparisons:
- Define `bool_comp_1` to be the logical `and` of `bool_1` and `bool_2`
- Define `bool_comp_2` to be the logical `or` of  `bool_1` and `bool_2`

```python
# These variables provided for you
bool_1 = True
bool_2 = False
```

<!-- grade_id: cell-bf6144adc1f1e7f7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7a05eded0276311c, points: 0.5, grade: true, solution: false -->
```python
## TESTS FOR Q4 [Note: Includes Hidden Tests]
assert isinstance(bool_comp_1, bool)

```

<!-- grade_id: cell-eeb7dcc8687131df, points: 0.5, grade: true, solution: false -->
```python
## TESTS FOR Q4 [Note: Includes Hidden Tests]
assert isinstance(bool_comp_2, bool)

```

### Q5 - Value Comparisons (1 point)

Do the following value comparisons:
- Check if the value stored in `comp_val_1` is less than or equal to the value stored in `comp_val_2`
    - Store the result of this comparison to variable called `comp_1`


- Check if the value stored in `comp_val_3` is not equal to (is different from) the value stored in `comp_val_4`
    - Store the result of this comparison to variable called `comp_2`

```python
# These variables provided for you
comp_val_1 = 13
comp_val_2 = 17
comp_val_3 = 23
comp_val_4 = 22
```

<!-- grade_id: cell-41f5d60aa2e09728, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-313f6320a6e8c70a, points: 0.5, grade: true, solution: false -->
```python
## TESTS FOR Q5 [Note: Includes Hidden Tests]
assert isinstance(comp_1, bool)

```

<!-- grade_id: cell-283d8c591109241d, points: 0.5, grade: true, solution: false -->
```python
## TESTS FOR Q5 [Note: Includes Hidden Tests]
assert isinstance(comp_2, bool)

```

## Part III - Conditionals

### Q6 - Conditionals With Booleans (1 point)

Write a conditional that tests a boolean variable `status`, specifically:
- if `status` evaluates as `True`, set the variable `output` as the string "GOOD"
- else (meaning `status` must evaluate as `False`), set the variable `output` as the string "BAD"

<!-- grade_id: cell-d2e812dc1f9a43a6, points: 0, grade: false, solution: true -->
```python
%%writefile A1Code/q6_code.py
# ^ You can ignore this line - it is used to help check your code

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-de368a5bcd2db9c1, points: 0.5, grade: true, solution: false -->
```python
## TESTS FOR Q6

# Set stats - check for True case
status = True

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q6_code.py  
assert output == "GOOD"
```

<!-- grade_id: cell-60a753f009ed6b1b, points: 0.5, grade: true, solution: false -->
```python
## TESTS FOR Q6

# Set stats - check for False case
status = False

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q6_code.py  
assert output == "BAD"
```

### Q7 - Conditionals With Comparisons (1 point)

Consider a case in which you are running an experiment, with multiple groups of subjects, including young subjects and old subjects. 

Based on the subjects age group, you want to do a different analysis - specifically if they are in the old group.

Write a conditional with an `if` and an `else`:
- The `if` should check if the variable `subj_age` has the value `'old'` (a string).
    - If so, it should set the variable `change_analysis` to the value `True` (a boolean)
    
    
- Else, the variable `change_analysis` should be set to False (a boolean)

<!-- grade_id: cell-5fdac280e980848a, points: 0, grade: false, solution: true -->
```python
%%writefile A1Code/q7_code.py
# ^ You can ignore this line - it is used to help check your code

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e59c9f37e0310a5, points: 0.5, grade: true, solution: false -->
```python
## TESTS FOR Q7

# Set subj_age to old
subj_age = 'old'

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q7_code.py  

# Check that code 
assert change_analysis
assert isinstance(change_analysis, bool)
```

<!-- grade_id: cell-c4621a021d61b90f, points: 0.5, grade: true, solution: false -->
```python
## TESTS FOR Q7

# Change subj_age
subj_age = 'young'

# Re-run code and make sure the conditional still works
%run -i ./A1Code/q7_code.py
assert not change_analysis
assert isinstance(change_analysis, bool)
```

### Q8 - Conditionals With Multiple Options (1 point)

Sometimes we may need to check multiple conditions, with many (more than two) possibilities.

Create a conditional with the following outline:
- If `check_1` and `check_2` variables are both True, it should set the value of a variable `outcome` to the string 'BOTH'
- Elif `check_1` is True and `check_2` is False, it should set the value of a variable `outcome` to the string 'ONE'
- Elif `check_1` is False and `check_2` is True, it should set the value of a variable `outcome` to the string 'TWO'
- Else (meaning both must be False), it should set the value of a variable `outcome` to the string 'NEITHER'

<!-- grade_id: cell-0cada564ef3bd740, points: 0, grade: false, solution: true -->
```python
%%writefile A1Code/q8_code.py
# ^ You can ignore this line - it is used to help check your code

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7b81b816f5fbd8fc, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q8

# Check 'BOTH' case
check_1 = True
check_2 = True

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q8_code.py
assert outcome == 'BOTH'
```

<!-- grade_id: cell-621a9631f8689eb8, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q8

# Check 'ONE' case
check_1 = True
check_2 = False

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q8_code.py
assert outcome == 'ONE'
```

<!-- grade_id: cell-a01f95735ce601a7, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q8

# Check 'TWO' case
check_1 = False
check_2 = True

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q8_code.py
assert outcome == 'TWO'
```

<!-- grade_id: cell-74ded29c5e0ab108, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q8

# Check 'NEITHER' case
check_1 = False
check_2 = False

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q8_code.py
assert outcome == 'NEITHER'
```

## Part IV - Putting it all together

### Q9 - Multiple Cases: Experiment Example (1 point)

Lets revisit the setup of Q7 - in which we wanted to set a variable to indicate if a subject in an experiment was in a particular group, and thus required different processing. In this question we will extend that premise, by writing a set of conditionals to mark an interaction between subject's age and score. 

For example, it is common for behavioural performance to decrease with age. However, some subjects do not seem to show this behavioural deficit with age. Investigating such 'super-ages' could be useful for learning about the cognitive decline of aging.

Write a conditional that does the following:
- If the variable `subj_age` has the value `old` and the variable `score` is greater than 25:
    - set the variable `subj_status` to the string 'Super Ager'
    
    
- Elif the variable `subj_age` has the value `old` and the variable `score` is less than or equal to 25:
    - set the variable `subj_status` to the string 'Normal Ager'
    
    
- Elif the variable `subj_age` has the value `young` and the variable `score` is greater than 25:
    - set the variable `subj_status` to the string 'Normal Youth'
    
    
- Elif the variable `subj_age` has the value `young` and the variable `score` is less than or equal to 25:
    - set the variable `subj_status` to the string 'Bad Youth'

<!-- grade_id: cell-514a6e20b819dee0, points: 0, grade: false, solution: true -->
```python
%%writefile A1Code/q9_code.py
# ^ You can ignore this line - it is used to help check your code

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-10e577dabe2f9a4f, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q9

subj_age = 'old'
score = 28

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q9_code.py
assert subj_status == 'Super Ager'
```

<!-- grade_id: cell-c259aac343bb297a, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q9

subj_age = 'old'
score = 22

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q9_code.py
assert subj_status == 'Normal Ager'
```

<!-- grade_id: cell-b05b4bec57b61e00, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q9

subj_age = 'young'
score = 29

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q9_code.py
assert subj_status == 'Normal Youth'
```

<!-- grade_id: cell-0cc09b0683fdfd57, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q9

subj_age = 'young'
score = 21

# This line is a way to run & re-run the code you wrote
%run -i ./A1Code/q9_code.py
assert subj_status == 'Bad Youth'
```

### Q10 - Creating a Calculator (1.5 points)

Finally we will write a flexible calculater program, that performs multiple operations (addition, subtraction, multiplication and division) on pairs of numbers. To do so, the code will use a string variable `operation` that indicates with operation to perform, and two numerical values, `num1` and `num2`, upon which the operation is performed. The calculator will assign the output of it's calculation to a variable called `answer`. 

To create this calculater, we will be using conditionals. 

In particular, create a calculator such that:
- if `operation` has the value 'add', then `answer` gets set as the sum of `num1` and `num2`
- elif `operation` has the value 'subtract', then `answer` gets set as the substraction of `num2` from `num2`
- elif `operation` has the value 'multiply', then `answer` gets set as the multiplication of `num1` and `num2`
- elif `operation` has the value 'divide', then `answer` gets set as the division of `num1` by `num2`

For an example, with the following variables to start with:

```
operation = 'multiply'
num1 = 2
num2 = 3
```
We should expect our program to execute and set the variable `answer` to be `6`.

<!-- grade_id: cell-d44d4aa343270410, points: 0, grade: false, solution: true -->
```python
%%writefile A1Code/q10_code.py

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Set numbers for tests
num1 = 21
num2 = 3
```

<!-- grade_id: cell-63991f35fb0fa2b7, points: 0.3, grade: true, solution: false -->
```python
## TESTS FOR Q10

# Test for sum
operation = 'add'
%run -i ./A1Code/q10_code.py
assert answer == 24
```

<!-- grade_id: cell-df854472375eaaea, points: 0.3, grade: true, solution: false -->
```python
## TESTS FOR Q10

# Test for subtract
operation = 'subtract'
%run -i ./A1Code/q10_code.py
assert answer == 18
```

<!-- grade_id: cell-5cb26332c714b7b6, points: 0.3, grade: true, solution: false -->
```python
## TESTS FOR Q10

# Test for multiply
operation = 'multiply'
%run -i ./A1Code/q10_code.py
assert answer == 63
```

<!-- grade_id: cell-515fa93f1045b72f, points: 0.3, grade: true, solution: false -->
```python
## TESTS FOR Q10

# Test for divide
operation = 'divide'
%run -i ./A1Code/q10_code.py
assert answer == 7
```

<!-- grade_id: cell-544aefbe440ec548, points: 0.3, grade: true, solution: false -->
```python
## TESTS FOR Q10

# Test unrecognized operation
operation = 'not a thing'
%run -i ./A1Code/q10_code.py
assert answer == None
```

## The End!

This is the end of the assignment!

Have a look back over your answers, and also make sure to `Restart & Run All` from the kernel menu to double check that everything is working properly. When you are ready to submit your assignment, validate and submit on DataHub!